<section xml:id="sec-HW9-1">
  <title>HW9</title>
  <introduction>
    <p>
    Main Concepts:
    <ul>
      <li>A text file is a sequence of lines, where each line is a string ending in the newline character, '\n'.</li>
      <li>Syntax and semantics for calling the <alert>open</alert>function to open a text file for reading.</li>
      <li>Syntax and semantics for some basic file methods: <alert>readline, readlines, read, close</alert>.</li>
      <li>Coding pattern: How to iterate over the lines of a text file using a for loop.</li>
    </ul>
    </p>
  </introduction>
  <exercises>
    <exercisegroup xml:id="exer-file_reading_match-1" label="exe-sec-HW9-1-225597">
      <introduction>
        <p>Suppose <alert>pet_names.txt</alert> is the name of a text file in our current directory, and we successfully opened the file for reading by executing the following statement:
        <cd>pets = open('pet_names.txt', 'r')</cd>
        </p>
        <p>Match the code chunk with the correct description of the result produced by executing the code.</p>
      </introduction>
      <exercise xml:id="exer-file_reading_match-0-1" label="exe-sec-HW9-1-225597-0">
        <statement>
          <pre>x = pets.readlines()  # NOTICE the 's' on readlines</pre>
        </statement>
        <choices multiple-correct="no" randomize="yes">
          <choice correct="yes">
            <statement>
              <p>x is a list of strings; each item in x is a line from pet_names.txt.</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is a single string containing the entire contents of pet_names.txt.</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the first line of pet_names.txt.</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the number of lines in pet_names.txt</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the number of characters in pet_names.txt, including newline characters</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the number of characters in pet_names.txt, NOT including newline characters</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the last line of pet_names.txt</p>
            </statement>
            <feedback />
          </choice>
        </choices>
      </exercise>
      <exercise xml:id="exer-file_reading_match-1-1" label="exe-sec-HW9-1-225597-1">
        <statement>
          <pre>x = pets.read()</pre>
        </statement>
        <choices multiple-correct="no" randomize="yes">
          <choice correct="no">
            <statement>
              <p>x is a list of strings; each item in x is a line from pet_names.txt.</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="yes">
            <statement>
              <p>x is a single string containing the entire contents of pet_names.txt.</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the first line of pet_names.txt.</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the number of lines in pet_names.txt</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the number of characters in pet_names.txt, including newline characters</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the number of characters in pet_names.txt, NOT including newline characters</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the last line of pet_names.txt</p>
            </statement>
            <feedback />
          </choice>
        </choices>
      </exercise>
      <exercise xml:id="exer-file_reading_match-2-1" label="exe-sec-HW9-1-225597-2">
        <statement>
          <pre>x = pets.readline()</pre>
        </statement>
        <choices multiple-correct="no" randomize="yes">
          <choice correct="no">
            <statement>
              <p>x is a list of strings; each item in x is a line from pet_names.txt.</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is a single string containing the entire contents of pet_names.txt.</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="yes">
            <statement>
              <p>x is the first line of pet_names.txt.</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the number of lines in pet_names.txt</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the number of characters in pet_names.txt, including newline characters</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the number of characters in pet_names.txt, NOT including newline characters</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the last line of pet_names.txt</p>
            </statement>
            <feedback />
          </choice>
        </choices>
      </exercise>
      <exercise xml:id="exer-file_reading_match-3-1" label="exe-sec-HW9-1-225597-3">
        <statement>
          <pre>x = len(pets.readlines())  # NOTICE the 's' on readlines</pre>
        </statement>
        <choices multiple-correct="no" randomize="yes">
          <choice correct="no">
            <statement>
              <p>x is a list of strings; each item in x is a line from pet_names.txt.</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is a single string containing the entire contents of pet_names.txt.</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the first line of pet_names.txt.</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="yes">
            <statement>
              <p>x is the number of lines in pet_names.txt</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the number of characters in pet_names.txt, including newline characters</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the number of characters in pet_names.txt, NOT including newline characters</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the last line of pet_names.txt</p>
            </statement>
            <feedback />
          </choice>
        </choices>
      </exercise>
      <exercise xml:id="exer-file_reading_match-4-1" label="exe-sec-HW9-1-225597-4">
        <statement>
          <pre>acc = 0
for line in pets:
    acc = acc + len(line)
x = acc
          </pre>
        </statement>
        <choices multiple-correct="no" randomize="yes">
          <choice correct="no">
            <statement>
              <p>x is a list of strings; each item in x is a line from pet_names.txt.</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is a single string containing the entire contents of pet_names.txt.</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the first line of pet_names.txt.</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the number of lines in pet_names.txt</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="yes">
            <statement>
              <p>x is the number of characters in pet_names.txt, including newline characters</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the number of characters in pet_names.txt, NOT including newline characters</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the last line of pet_names.txt</p>
            </statement>
            <feedback />
          </choice>
        </choices>
      </exercise>
      <exercise xml:id="exer-file_reading_match-5-1" label="exe-sec-HW9-1-225597-5">
        <statement>
          <pre>for line in pets:
    print("hello")
x = line
          </pre>
        </statement>
        <choices multiple-correct="no" randomize="yes">
          <choice correct="no">
            <statement>
              <p>x is a list of strings; each item in x is a line from pet_names.txt.</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is a single string containing the entire contents of pet_names.txt.</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the first line of pet_names.txt.</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the number of lines in pet_names.txt</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the number of characters in pet_names.txt, including newline characters</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="no">
            <statement>
              <p>x is the number of characters in pet_names.txt, NOT including newline characters</p>
            </statement>
            <feedback />
          </choice>
          <choice correct="yes">
            <statement>
              <p>x is the last line of pet_names.txt</p>
            </statement>
            <feedback />
          </choice>
        </choices>
      </exercise>
    </exercisegroup>
    <exercise xml:id="exer-print_first_line-1" label="exe-sec-HW9-1-32177">
      <title>Print First Line</title>
      <statement>
        <p>The function <c>print_first_line</c> takes one parameter, <c>fname</c>, the name of a text file. The function should open the file for reading, print the first line from the file, and close the file.</p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="60%" />
          <col top="minor" width="40%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>print_first_line("wordlist1.txt")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>happy</cd>
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <preamble visible="no">
_datafiles = {'cards.txt': 'Three of Clubs\nSix of Diamonds\nTen of Hearts\nEight of Clubs\nTwo of Clubs\nSix of Spades\nEight of Spades\nThree of Diamonds\nQueen of Diamonds\nTwo of Spades\nKing of Spades\nFive of Spades\nJack of Clubs\nFive of Hearts\nFour of Clubs\nNine of Clubs\nFive of Clubs\nTen of Spades\nThree of Spades\nJack of Diamonds\nNine of Spades\nSeven of Hearts\nEight of Diamonds\nSeven of Clubs\nFive of Diamonds\nQueen of Clubs\nJack of Hearts\nKing of Diamonds\nSix of Hearts\nFour of Diamonds\nAce of Hearts\nQueen of Hearts\nTwo of Diamonds\nKing of Clubs\nTwo of Hearts\nAce of Spades\nFour of Hearts\nFour of Spades\nNine of Diamonds\nSeven of Spades\nEight of Hearts\nSeven of Diamonds\nKing of Hearts\nSix of Clubs\nAce of Clubs\nNine of Hearts\nJack of Spades\nTen of Diamonds\nQueen of Spades\nThree of Hearts\nTen of Clubs\nAce of Diamonds\n', 'clocktimes.txt': '09:05\n09:15\n09:20\n09:30\n10:20\n10:25\n10:30\n10:40\n10:45\n11:00\n11:15\n11:30\n11:35\n11:55\n12:00\n12:05\n12:10\n12:45\n12:50\n12:55\n13:05\n13:20\n13:35\n14:00\n14:10\n14:15\n14:20\n14:25\n14:35\n14:45\n14:50\n15:00\n15:55\n16:00\n16:05\n16:20\n16:25\n16:30\n16:35\n16:40\n16:50\n16:55\n', 'jazzscales.txt': 'Play a major pent scale in Ab with attack on the 3rd string\nPlay a 7th pent scale in C# (Db) with attack on the 1st string\nPlay a major scale in B with attack on the 1st string\n', 'random1.txt': '77635\n51437\n25094\n13874\n24429\n88575\n90277\n72340\n8826\n36166\n99144\n6920\n56237\n58141\n77661\n51028\n93722\n14354\n26125\n10151\n75322\n', 'strings1.txt': '43XkyLl3QO\nKgYH4yZ9vp\nB9z0d27BWf\n3lVimQBNOl\nPZF1FkRCC0\naodmjRnbSM\nQinMgOKyJ3\nCIRpxzuBov\nbl2VB97vGe\nypinbN0yjw\nJotIY4A9GC\nwjlureOXLy\nDlHXQlUuhv\nQHvicuwzKp\nMzORTv2BYU\nvVnuElBS0s\naezB6hMOqj\nhipsag5OAM\nLhgNsYqWhB\n9b3yi6e29s\n6fx3eADW4n\nV28c1LIoQC\nIsN0xS6bZ4\nd1e0b2hJX6\nRwXQZcECRu\n', 'wordlist1.txt': 'happy\nsad\nangsty\neuphoric\nmaudlin\n'}

# Start of FileLike
class FileLike:
  def __init__(self, contents):
    self.contents = contents
    self.cursor = 0

  def __enter__(self):
    return self

  def __exit__(self, exc_type, exc_value, traceback):
    self.close()

  def tell(self):
    return self.cursor

  def seek(self, n):
    self.cursor = n

  def read(self):
    n = self.cursor
    self.cursor = max(n, len(self.contents))
    return self.contents[n:]

  def readline(self):
    n = self.cursor
    rest = self.contents[n:]
    index = rest.find("\n")
    if index == -1:
      self.cursor = len(self.contents)
      return rest
    index += 1
    self.cursor = n + index
    return rest[:index]

  def readlines(self):
    results = []
    while True:
      line = self.readline()
      if line == '':
        return results
      results.append(line)

  def __iter__(self):
    for l in self.readlines():
      yield l


  def close(self):
    pass

# End of FileLike

def open(fname, mode = "r"):
  """our own open"""
  if fname in _datafiles:
    return FileLike(_datafiles[fname])
  raise BaseException(f"File not found: {fname}")

        </preamble>
        <code>

        </code>
<tests>
from unittest.gui import TestCaseGui
import sys
import StringIO

class myTests(TestCaseGui):

  def test1(self):
    sys.stdout = StringIO.StringIO()
    print_first_line('cards.txt')
    printout = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    self.assertEqual(printout, 'Three of Clubs\n\n', "print_first_line('cards.txt')")

  def test2(self):
    sys.stdout = StringIO.StringIO()
    print_first_line('clocktimes.txt')
    printout = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    self.assertEqual(printout, '09:05\n\n', "print_first_line('clocktimes.txt')")

  def test3(self):
    sys.stdout = StringIO.StringIO()
    print_first_line('jazzscales.txt')
    printout = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    self.assertEqual(printout, 'Play a major pent scale in Ab with attack on the 3rd string\n\n', "print_first_line('jazzscales.txt')")

  def test4(self):
    sys.stdout = StringIO.StringIO()
    print_first_line('random1.txt')
    printout = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    self.assertEqual(printout, '77635\n\n', "print_first_line('random1.txt')")

  def test5(self):
    sys.stdout = StringIO.StringIO()
    print_first_line('strings1.txt')
    printout = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    self.assertEqual(printout, '43XkyLl3QO\n\n', "print_first_line('strings1.txt')")

  def test6(self):
    sys.stdout = StringIO.StringIO()
    print_first_line('wordlist1.txt')
    printout = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    self.assertEqual(printout, 'happy\n\n', "print_first_line('wordlist1.txt')")

myTests().main()
</tests>
      </program>
    </exercise>
    <exercise xml:id="exer-file_get_second_line-1" label="exe-sec-HW9-1-32165">
      <title>Get Second Line</title>
      <statement>
        <p>The function <c>get_second_line</c> takes one parameter, <c>fname</c>, the name of a text file. The function should open the file for reading, determine the <alert>second</alert>line of the file, close the file, and return the second line (a string).</p>
        <p>NOTE: You may assume the file will have at least two lines.</p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="60%" />
          <col top="minor" width="40%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>get_second_line("cards.txt")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>Six of Diamonds</cd>
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <preamble visible="no">
_datafiles = {'cards.txt': 'Three of Clubs\nSix of Diamonds\nTen of Hearts\nEight of Clubs\nTwo of Clubs\nSix of Spades\nEight of Spades\nThree of Diamonds\nQueen of Diamonds\nTwo of Spades\nKing of Spades\nFive of Spades\nJack of Clubs\nFive of Hearts\nFour of Clubs\nNine of Clubs\nFive of Clubs\nTen of Spades\nThree of Spades\nJack of Diamonds\nNine of Spades\nSeven of Hearts\nEight of Diamonds\nSeven of Clubs\nFive of Diamonds\nQueen of Clubs\nJack of Hearts\nKing of Diamonds\nSix of Hearts\nFour of Diamonds\nAce of Hearts\nQueen of Hearts\nTwo of Diamonds\nKing of Clubs\nTwo of Hearts\nAce of Spades\nFour of Hearts\nFour of Spades\nNine of Diamonds\nSeven of Spades\nEight of Hearts\nSeven of Diamonds\nKing of Hearts\nSix of Clubs\nAce of Clubs\nNine of Hearts\nJack of Spades\nTen of Diamonds\nQueen of Spades\nThree of Hearts\nTen of Clubs\nAce of Diamonds\n', 'clocktimes.txt': '09:05\n09:15\n09:20\n09:30\n10:20\n10:25\n10:30\n10:40\n10:45\n11:00\n11:15\n11:30\n11:35\n11:55\n12:00\n12:05\n12:10\n12:45\n12:50\n12:55\n13:05\n13:20\n13:35\n14:00\n14:10\n14:15\n14:20\n14:25\n14:35\n14:45\n14:50\n15:00\n15:55\n16:00\n16:05\n16:20\n16:25\n16:30\n16:35\n16:40\n16:50\n16:55\n', 'jazzscales.txt': 'Play a major pent scale in Ab with attack on the 3rd string\nPlay a 7th pent scale in C# (Db) with attack on the 1st string\nPlay a major scale in B with attack on the 1st string\n', 'random1.txt': '77635\n51437\n25094\n13874\n24429\n88575\n90277\n72340\n8826\n36166\n99144\n6920\n56237\n58141\n77661\n51028\n93722\n14354\n26125\n10151\n75322\n', 'strings1.txt': '43XkyLl3QO\nKgYH4yZ9vp\nB9z0d27BWf\n3lVimQBNOl\nPZF1FkRCC0\naodmjRnbSM\nQinMgOKyJ3\nCIRpxzuBov\nbl2VB97vGe\nypinbN0yjw\nJotIY4A9GC\nwjlureOXLy\nDlHXQlUuhv\nQHvicuwzKp\nMzORTv2BYU\nvVnuElBS0s\naezB6hMOqj\nhipsag5OAM\nLhgNsYqWhB\n9b3yi6e29s\n6fx3eADW4n\nV28c1LIoQC\nIsN0xS6bZ4\nd1e0b2hJX6\nRwXQZcECRu\n', 'wordlist1.txt': 'happy\nsad\nangsty\neuphoric\nmaudlin\n'}

# Start of FileLike
class FileLike:
  def __init__(self, contents):
    self.contents = contents
    self.cursor = 0

  def __enter__(self):
    return self

  def __exit__(self, exc_type, exc_value, traceback):
    self.close()

  def tell(self):
    return self.cursor

  def seek(self, n):
    self.cursor = n

  def read(self):
    n = self.cursor
    self.cursor = max(n, len(self.contents))
    return self.contents[n:]

  def readline(self):
    n = self.cursor
    rest = self.contents[n:]
    index = rest.find("\n")
    if index == -1:
      self.cursor = len(self.contents)
      return rest
    index += 1
    self.cursor = n + index
    return rest[:index]

  def readlines(self):
    results = []
    while True:
      line = self.readline()
      if line == '':
        return results
      results.append(line)

  def __iter__(self):
    for l in self.readlines():
      yield l


  def close(self):
    pass

# End of FileLike

def open(fname, mode = "r"):
  """our own open"""
  if fname in _datafiles:
    return FileLike(_datafiles[fname])
  raise BaseException(f"File not found: {fname}")

        </preamble>
        <code>

        </code>
<tests>
from unittest.gui import TestCaseGui
import StringIO
import sys

class myTests(TestCaseGui):


  def test1(self):
    res = get_second_line('cards.txt')
    self.assertEqual(res, 'Six of Diamonds\n', "get_second_line('cards.txt')")

  def test2(self):
    res = get_second_line('clocktimes.txt')
    self.assertEqual(res, '09:15\n', "get_second_line('clocktimes.txt')")

  def test3(self):
    res = get_second_line('jazzscales.txt')
    self.assertEqual(res, 'Play a 7th pent scale in C# (Db) with attack on the 1st string\n', "get_second_line('jazzscales.txt')")

  def test4(self):
    res = get_second_line('random1.txt')
    self.assertEqual(res, '51437\n', "get_second_line('random1.txt')")

  def test5(self):
    res = get_second_line('strings1.txt')
    self.assertEqual(res, 'KgYH4yZ9vp\n', "get_second_line('strings1.txt')")

  def test6(self):
    res = get_second_line('wordlist1.txt')
    self.assertEqual(res, 'sad\n', "get_second_line('wordlist1.txt')")

myTests().main()
</tests>
      </program>
    </exercise>
    <exercise xml:id="exer-file_iteration_introduction-1" label="exe-sec-HW9-1-271603">
      <title>File Iteration: Introduction</title>
      <statement>
        <p>A text file is essentially a sequence of lines. Here is a <alert>typical coding pattern</alert> for processing each line of a file in turn:</p>
        <p>
          <ul>
            <li>Open the file for reading.</li>
            <li>Use a <alert>for loop</alert> to iterate over the lines of text, executing a block of one or more statements to process each line.</li>
            <li>Close the file.</li>
            <li>Finish processing (print, return, etc.) if needed</li>
          </ul>
        </p>
        <p>NOTE: You will need to use this pattern, or some variation of it, to solve the remaining coding problems in this assignment.</p>
        <p>As the following for loop iterates through each line of the file, the loop variable (<c>line</c>) represents the current line of the file as a string of characters.</p>
        <pre>myFile = open("someTextFile.txt", "r")
for line in myFile:
    do something on each iteration
    do a second thing on each iteration
    etc...
myFile.close()
        </pre>
        <p><alert>PROBLEM</alert>: Read a string from standard input. Then, iterate over a text file called "cards.txt" and, for each line of the file that starts with the string read from standard input, print that line.</p>
        <p>
          <alert>TO DO:</alert>
          <ul>
            <li>Add the missing bits of code to complete the partial solution that is provided (replace the _________ s).</li>
            <li><em>Follow the coding pattern described above.</em></li>
          </ul>
        </p>
        <p><alert>NOTE</alert>: If there are no lines in the file beginning with the given string, there should be no output.</p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="30%" />
          <col top="minor" width="70%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Input</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>Ace</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>Ace of Hearts

Ace of Spades

Ace of Clubs

Ace of Diamonds</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>5</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd />
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <preamble visible="no">
_datafiles = {'cards.txt': 'Three of Clubs\nSix of Diamonds\nTen of Hearts\nEight of Clubs\nTwo of Clubs\nSix of Spades\nEight of Spades\nThree of Diamonds\nQueen of Diamonds\nTwo of Spades\nKing of Spades\nFive of Spades\nJack of Clubs\nFive of Hearts\nFour of Clubs\nNine of Clubs\nFive of Clubs\nTen of Spades\nThree of Spades\nJack of Diamonds\nNine of Spades\nSeven of Hearts\nEight of Diamonds\nSeven of Clubs\nFive of Diamonds\nQueen of Clubs\nJack of Hearts\nKing of Diamonds\nSix of Hearts\nFour of Diamonds\nAce of Hearts\nQueen of Hearts\nTwo of Diamonds\nKing of Clubs\nTwo of Hearts\nAce of Spades\nFour of Hearts\nFour of Spades\nNine of Diamonds\nSeven of Spades\nEight of Hearts\nSeven of Diamonds\nKing of Hearts\nSix of Clubs\nAce of Clubs\nNine of Hearts\nJack of Spades\nTen of Diamonds\nQueen of Spades\nThree of Hearts\nTen of Clubs\nAce of Diamonds\n', 'clocktimes.txt': '09:05\n09:15\n09:20\n09:30\n10:20\n10:25\n10:30\n10:40\n10:45\n11:00\n11:15\n11:30\n11:35\n11:55\n12:00\n12:05\n12:10\n12:45\n12:50\n12:55\n13:05\n13:20\n13:35\n14:00\n14:10\n14:15\n14:20\n14:25\n14:35\n14:45\n14:50\n15:00\n15:55\n16:00\n16:05\n16:20\n16:25\n16:30\n16:35\n16:40\n16:50\n16:55\n', 'jazzscales.txt': 'Play a major pent scale in Ab with attack on the 3rd string\nPlay a 7th pent scale in C# (Db) with attack on the 1st string\nPlay a major scale in B with attack on the 1st string\n', 'random1.txt': '77635\n51437\n25094\n13874\n24429\n88575\n90277\n72340\n8826\n36166\n99144\n6920\n56237\n58141\n77661\n51028\n93722\n14354\n26125\n10151\n75322\n', 'strings1.txt': '43XkyLl3QO\nKgYH4yZ9vp\nB9z0d27BWf\n3lVimQBNOl\nPZF1FkRCC0\naodmjRnbSM\nQinMgOKyJ3\nCIRpxzuBov\nbl2VB97vGe\nypinbN0yjw\nJotIY4A9GC\nwjlureOXLy\nDlHXQlUuhv\nQHvicuwzKp\nMzORTv2BYU\nvVnuElBS0s\naezB6hMOqj\nhipsag5OAM\nLhgNsYqWhB\n9b3yi6e29s\n6fx3eADW4n\nV28c1LIoQC\nIsN0xS6bZ4\nd1e0b2hJX6\nRwXQZcECRu\n', 'wordlist1.txt': 'happy\nsad\nangsty\neuphoric\nmaudlin\n'}

import StringIO
import sys

# Start of FileLike
class FileLike:
  def __init__(self, contents):
    self.contents = contents
    self.cursor = 0

  def __enter__(self):
    return self

  def __exit__(self, exc_type, exc_value, traceback):
    self.close()

  def tell(self):
    return self.cursor

  def seek(self, n):
    self.cursor = n

  def read(self):
    n = self.cursor
    self.cursor = max(n, len(self.contents))
    return self.contents[n:]

  def readline(self):
    n = self.cursor
    rest = self.contents[n:]
    index = rest.find("\n")
    if index == -1:
      self.cursor = len(self.contents)
      return rest
    index += 1
    self.cursor = n + index
    return rest[:index]

  def readlines(self):
    results = []
    while True:
      line = self.readline()
      if line == '':
        return results
      results.append(line)

  def __iter__(self):
    for l in self.readlines():
      yield l

  def close(self):
    pass

# End of FileLike

def open(fname, mode = "r"):
  """our own open"""
  if fname in _datafiles:
    return FileLike(_datafiles[fname])
  raise BaseException(f"File not found: {fname}")

sys.stdin = StringIO.StringIO("Ace")
sys.stdout = StringIO.StringIO()
        </preamble>
        <code>
# read the string from input
matchString = ____________

# open the file and assign the file object to a variable
myFile = _____________

# iterate over the lines of the file, watching for lines
# which begin with matchString
for line in _____________
    if line[0:len(matchString)] == matchString:
        # print here
        __________________
myFile.close()
        </code>
        <postamble visible="no">
sys.stdin = sys.__stdin__
sys.stdout = sys.__stdout__
        </postamble>
<tests>
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):


  def test1(self):
    sys.stdin = StringIO.StringIO('Ace')
    sys.stdout = StringIO.StringIO()
    exec(self.getEditorText())
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    sys.stdin = sys.__stdin__
    self.assertEqual(res, 'Ace of Hearts\n\nAce of Spades\n\nAce of Clubs\n\nAce of Diamonds\n\n', "Input: Ace")

  def test2(self):
    sys.stdin = StringIO.StringIO('Jack')
    sys.stdout = StringIO.StringIO()
    exec(self.getEditorText())
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    sys.stdin = sys.__stdin__
    self.assertEqual(res, 'Jack of Clubs\n\nJack of Diamonds\n\nJack of Hearts\n\nJack of Spades\n\n', "Input: Jack")

  def test3(self):
    sys.stdin = StringIO.StringIO('queen')
    sys.stdout = StringIO.StringIO()
    exec(self.getEditorText())
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    sys.stdin = sys.__stdin__
    self.assertEqual(res, '', "Input: queen")

  def test4(self):
    sys.stdin = StringIO.StringIO('Five')
    sys.stdout = StringIO.StringIO()
    exec(self.getEditorText())
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    sys.stdin = sys.__stdin__
    self.assertEqual(res, 'Five of Spades\n\nFive of Hearts\n\nFive of Clubs\n\nFive of Diamonds\n\n', "Input: Five")

  def test5(self):
    sys.stdin = StringIO.StringIO('5')
    sys.stdout = StringIO.StringIO()
    exec(self.getEditorText())
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    sys.stdin = sys.__stdin__
    self.assertEqual(res, '', "Input: 5")

  def test6(self):
    sys.stdin = StringIO.StringIO('Nine')
    sys.stdout = StringIO.StringIO()
    exec(self.getEditorText())
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    sys.stdin = sys.__stdin__
    self.assertEqual(res, 'Nine of Clubs\n\nNine of Spades\n\nNine of Diamonds\n\nNine of Hearts\n\n', "Input: Nine")

  def test7(self):
    sys.stdin = StringIO.StringIO('T')
    sys.stdout = StringIO.StringIO()
    exec(self.getEditorText())
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    sys.stdin = sys.__stdin__
    self.assertEqual(res, 'Three of Clubs\n\nTen of Hearts\n\nTwo of Clubs\n\nThree of Diamonds\n\nTwo of Spades\n\nTen of Spades\n\nThree of Spades\n\nTwo of Diamonds\n\nTwo of Hearts\n\nTen of Diamonds\n\nThree of Hearts\n\nTen of Clubs\n\n', "Input: T")

myTests().main()
</tests>

      </program>
    </exercise>
    <exercise xml:id="exer-file_first_chars-1" label="exe-sec-HW9-1-32176">
      <title>File First Chars</title>
      <statement>
        <p>The function <c>file_first_chars</c> takes one parameter, <c>fname</c>, the name of a text file, and <alert>returns</alert>a string made up of the first character from each line in the file.  You may assume there are no blank lines in the file.</p>
        <p><alert>Hint</alert>: Use a <alert>for</alert>loop to iterate over the lines of the file and accumulate the string of first characters.</p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="70%" />
          <col top="minor" width="30%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>file_first_chars("wordlist1.txt")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>hsaem</cd>
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <preamble visible="no">
_datafiles = {'cards.txt': 'Three of Clubs\nSix of Diamonds\nTen of Hearts\nEight of Clubs\nTwo of Clubs\nSix of Spades\nEight of Spades\nThree of Diamonds\nQueen of Diamonds\nTwo of Spades\nKing of Spades\nFive of Spades\nJack of Clubs\nFive of Hearts\nFour of Clubs\nNine of Clubs\nFive of Clubs\nTen of Spades\nThree of Spades\nJack of Diamonds\nNine of Spades\nSeven of Hearts\nEight of Diamonds\nSeven of Clubs\nFive of Diamonds\nQueen of Clubs\nJack of Hearts\nKing of Diamonds\nSix of Hearts\nFour of Diamonds\nAce of Hearts\nQueen of Hearts\nTwo of Diamonds\nKing of Clubs\nTwo of Hearts\nAce of Spades\nFour of Hearts\nFour of Spades\nNine of Diamonds\nSeven of Spades\nEight of Hearts\nSeven of Diamonds\nKing of Hearts\nSix of Clubs\nAce of Clubs\nNine of Hearts\nJack of Spades\nTen of Diamonds\nQueen of Spades\nThree of Hearts\nTen of Clubs\nAce of Diamonds\n', 'clocktimes.txt': '09:05\n09:15\n09:20\n09:30\n10:20\n10:25\n10:30\n10:40\n10:45\n11:00\n11:15\n11:30\n11:35\n11:55\n12:00\n12:05\n12:10\n12:45\n12:50\n12:55\n13:05\n13:20\n13:35\n14:00\n14:10\n14:15\n14:20\n14:25\n14:35\n14:45\n14:50\n15:00\n15:55\n16:00\n16:05\n16:20\n16:25\n16:30\n16:35\n16:40\n16:50\n16:55\n', 'jazzscales.txt': 'Play a major pent scale in Ab with attack on the 3rd string\nPlay a 7th pent scale in C# (Db) with attack on the 1st string\nPlay a major scale in B with attack on the 1st string\n', 'random1.txt': '77635\n51437\n25094\n13874\n24429\n88575\n90277\n72340\n8826\n36166\n99144\n6920\n56237\n58141\n77661\n51028\n93722\n14354\n26125\n10151\n75322\n', 'strings1.txt': '43XkyLl3QO\nKgYH4yZ9vp\nB9z0d27BWf\n3lVimQBNOl\nPZF1FkRCC0\naodmjRnbSM\nQinMgOKyJ3\nCIRpxzuBov\nbl2VB97vGe\nypinbN0yjw\nJotIY4A9GC\nwjlureOXLy\nDlHXQlUuhv\nQHvicuwzKp\nMzORTv2BYU\nvVnuElBS0s\naezB6hMOqj\nhipsag5OAM\nLhgNsYqWhB\n9b3yi6e29s\n6fx3eADW4n\nV28c1LIoQC\nIsN0xS6bZ4\nd1e0b2hJX6\nRwXQZcECRu\n', 'wordlist1.txt': 'happy\nsad\nangsty\neuphoric\nmaudlin\n'}

# Start of FileLike
class FileLike:
  def __init__(self, contents):
    self.contents = contents
    self.cursor = 0

  def __enter__(self):
    return self

  def __exit__(self, exc_type, exc_value, traceback):
    self.close()

  def tell(self):
    return self.cursor

  def seek(self, n):
    self.cursor = n

  def read(self):
    n = self.cursor
    self.cursor = max(n, len(self.contents))
    return self.contents[n:]

  def readline(self):
    n = self.cursor
    rest = self.contents[n:]
    index = rest.find("\n")
    if index == -1:
      self.cursor = len(self.contents)
      return rest
    index += 1
    self.cursor = n + index
    return rest[:index]

  def readlines(self):
    results = []
    while True:
      line = self.readline()
      if line == '':
        return results
      results.append(line)

  def __iter__(self):
    for l in self.readlines():
      yield l


  def close(self):
    pass

# End of FileLike

def open(fname, mode = "r"):
  """our own open"""
  if fname in _datafiles:
    return FileLike(_datafiles[fname])
  raise BaseException(f"File not found: {fname}")

        </preamble>
        <code>

        </code>
<tests>
from unittest.gui import TestCaseGui
import StringIO
import sys

class myTests(TestCaseGui):


  def test1(self):
    res = file_first_chars('cards.txt')
    self.assertEqual(res, 'TSTETSETQTKFJFFNFTTJNSESFQJKSFAQTKTAFFNSESKSANJTQTTA', "file_first_chars('cards.txt')")

  def test2(self):
    res = file_first_chars('clocktimes.txt')
    self.assertEqual(res, '000011111111111111111111111111111111111111', "file_first_chars('clocktimes.txt')")

  def test3(self):
    res = file_first_chars('jazzscales.txt')
    self.assertEqual(res, 'PPP', "file_first_chars('jazzscales.txt')")

  def test4(self):
    res = file_first_chars('random1.txt')
    self.assertEqual(res, '752128978396557591217', "file_first_chars('random1.txt')")

  def test5(self):
    res = file_first_chars('strings1.txt')
    self.assertEqual(res, '4KB3PaQCbyJwDQMvahL96VIdR', "file_first_chars('strings1.txt')")

  def test6(self):
    res = file_first_chars('wordlist1.txt')
    self.assertEqual(res, 'hsaem', "file_first_chars('wordlist1.txt')")

myTests().main()
</tests>
      </program>
    </exercise>
    <exercise xml:id="exer-file_first_words-1" label="exe-sec-HW9-1-32145">
      <title>File First Words</title>
      <statement>
        <p>The function <c>first_words</c> takes one parameter, <c>fname</c>, the name of a text file, and returns a list containing the first word from each line of the file. For example, if the file contents are:</p>
        <pre>apples are red
bananas are yellow
limes are green
        </pre>
        <p>then the list <c>["apples", "bananas", "limes"]</c> should be returned.NOTE: You may assume the file will contain no blank lines.BIG HINT: If <c>line</c>is a string representing a line of text (inside of a for loop!), then <c>L = line.split()</c> creates a list of the words in the line.</p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="40%" />
          <col top="minor" width="60%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>first_words("snakes.txt")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>['Cottonmouth', 'Timber', 'Black', 'Tiger', 'Copperhead',
 'Eastern', 'Western', 'Eastern', 'Prairie', 'Mojave']</cd>
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <preamble visible="no">
_datafiles = {'cards.txt': 'Three of Clubs\nSix of Diamonds\nTen of Hearts\nEight of Clubs\nTwo of Clubs\nSix of Spades\nEight of Spades\nThree of Diamonds\nQueen of Diamonds\nTwo of Spades\nKing of Spades\nFive of Spades\nJack of Clubs\nFive of Hearts\nFour of Clubs\nNine of Clubs\nFive of Clubs\nTen of Spades\nThree of Spades\nJack of Diamonds\nNine of Spades\nSeven of Hearts\nEight of Diamonds\nSeven of Clubs\nFive of Diamonds\nQueen of Clubs\nJack of Hearts\nKing of Diamonds\nSix of Hearts\nFour of Diamonds\nAce of Hearts\nQueen of Hearts\nTwo of Diamonds\nKing of Clubs\nTwo of Hearts\nAce of Spades\nFour of Hearts\nFour of Spades\nNine of Diamonds\nSeven of Spades\nEight of Hearts\nSeven of Diamonds\nKing of Hearts\nSix of Clubs\nAce of Clubs\nNine of Hearts\nJack of Spades\nTen of Diamonds\nQueen of Spades\nThree of Hearts\nTen of Clubs\nAce of Diamonds\n', 'clocktimes.txt': '09:05\n09:15\n09:20\n09:30\n10:20\n10:25\n10:30\n10:40\n10:45\n11:00\n11:15\n11:30\n11:35\n11:55\n12:00\n12:05\n12:10\n12:45\n12:50\n12:55\n13:05\n13:20\n13:35\n14:00\n14:10\n14:15\n14:20\n14:25\n14:35\n14:45\n14:50\n15:00\n15:55\n16:00\n16:05\n16:20\n16:25\n16:30\n16:35\n16:40\n16:50\n16:55\n', 'jazzscales.txt': 'Play a major pent scale in Ab with attack on the 3rd string\nPlay a 7th pent scale in C# (Db) with attack on the 1st string\nPlay a major scale in B with attack on the 1st string\n', 'random1.txt': '77635\n51437\n25094\n13874\n24429\n88575\n90277\n72340\n8826\n36166\n99144\n6920\n56237\n58141\n77661\n51028\n93722\n14354\n26125\n10151\n75322\n', 'strings1.txt': '43XkyLl3QO\nKgYH4yZ9vp\nB9z0d27BWf\n3lVimQBNOl\nPZF1FkRCC0\naodmjRnbSM\nQinMgOKyJ3\nCIRpxzuBov\nbl2VB97vGe\nypinbN0yjw\nJotIY4A9GC\nwjlureOXLy\nDlHXQlUuhv\nQHvicuwzKp\nMzORTv2BYU\nvVnuElBS0s\naezB6hMOqj\nhipsag5OAM\nLhgNsYqWhB\n9b3yi6e29s\n6fx3eADW4n\nV28c1LIoQC\nIsN0xS6bZ4\nd1e0b2hJX6\nRwXQZcECRu\n', 'wordlist1.txt': 'happy\nsad\nangsty\neuphoric\nmaudlin\n', 'snakes.txt': 'Cottonmouth (Agkistrodon piscivorus)\r\nTimber rattlesnake (Crotalus horridus)\r\nBlack diamond rattlesnake (Crotalus oreganus)\r\nTiger rattlesnake (Crotalus tigris)\r\nCopperhead (Agkistrodon contortrix)\r\nEastern coral snake (Micrurus fulvius)\r\nWestern diamondback rattlesnake (Crotalus atrox)\r\nEastern diamondback rattlesnake (Crotalus adamanteus)\r\nPrairie rattlesnake (Crotalus viridis)\r\nMojave rattlesnake (Crotalus scutulatus)'}

# Start of FileLike
class FileLike:
  def __init__(self, contents):
    self.contents = contents
    self.cursor = 0

  def __enter__(self):
    return self

  def __exit__(self, exc_type, exc_value, traceback):
    self.close()

  def tell(self):
    return self.cursor

  def seek(self, n):
    self.cursor = n

  def read(self):
    n = self.cursor
    self.cursor = max(n, len(self.contents))
    return self.contents[n:]

  def readline(self):
    n = self.cursor
    rest = self.contents[n:]
    index = rest.find("\n")
    if index == -1:
      self.cursor = len(self.contents)
      return rest
    index += 1
    self.cursor = n + index
    return rest[:index]

  def readlines(self):
    results = []
    while True:
      line = self.readline()
      if line == '':
        return results
      results.append(line)

  def __iter__(self):
    for l in self.readlines():
      yield l


  def close(self):
    pass

# End of FileLike

def open(fname, mode = "r"):
  """our own open"""
  if fname in _datafiles:
    return FileLike(_datafiles[fname])
  raise BaseException(f"File not found: {fname}")

        </preamble>
        <code>

        </code>
<tests>
from unittest.gui import TestCaseGui
import StringIO
import sys

class myTests(TestCaseGui):


  def test1(self):
    res = first_words('cards.txt')
    self.assertEqual(res, ['Three', 'Six', 'Ten', 'Eight', 'Two', 'Six', 'Eight', 'Three', 'Queen', 'Two', 'King', 'Five', 'Jack', 'Five', 'Four', 'Nine', 'Five', 'Ten', 'Three', 'Jack', 'Nine', 'Seven', 'Eight', 'Seven', 'Five', 'Queen', 'Jack', 'King', 'Six', 'Four', 'Ace', 'Queen', 'Two', 'King', 'Two', 'Ace', 'Four', 'Four', 'Nine', 'Seven', 'Eight', 'Seven', 'King', 'Six', 'Ace', 'Nine', 'Jack', 'Ten', 'Queen', 'Three', 'Ten', 'Ace'], "first_words('cards.txt')")

  def test2(self):
    res = first_words('clocktimes.txt')
    self.assertEqual(res, ['09:05', '09:15', '09:20', '09:30', '10:20', '10:25', '10:30', '10:40', '10:45', '11:00', '11:15', '11:30', '11:35', '11:55', '12:00', '12:05', '12:10', '12:45', '12:50', '12:55', '13:05', '13:20', '13:35', '14:00', '14:10', '14:15', '14:20', '14:25', '14:35', '14:45', '14:50', '15:00', '15:55', '16:00', '16:05', '16:20', '16:25', '16:30', '16:35', '16:40', '16:50', '16:55'], "first_words('clocktimes.txt')")

  def test3(self):
    res = first_words('jazzscales.txt')
    self.assertEqual(res, ['Play', 'Play', 'Play'], "first_words('jazzscales.txt')")

  def test4(self):
    res = first_words('random1.txt')
    self.assertEqual(res, ['77635', '51437', '25094', '13874', '24429', '88575', '90277', '72340', '8826', '36166', '99144', '6920', '56237', '58141', '77661', '51028', '93722', '14354', '26125', '10151', '75322'], "first_words('random1.txt')")

  def test5(self):
    res = first_words('strings1.txt')
    self.assertEqual(res, ['43XkyLl3QO', 'KgYH4yZ9vp', 'B9z0d27BWf', '3lVimQBNOl', 'PZF1FkRCC0', 'aodmjRnbSM', 'QinMgOKyJ3', 'CIRpxzuBov', 'bl2VB97vGe', 'ypinbN0yjw', 'JotIY4A9GC', 'wjlureOXLy', 'DlHXQlUuhv', 'QHvicuwzKp', 'MzORTv2BYU', 'vVnuElBS0s', 'aezB6hMOqj', 'hipsag5OAM', 'LhgNsYqWhB', '9b3yi6e29s', '6fx3eADW4n', 'V28c1LIoQC', 'IsN0xS6bZ4', 'd1e0b2hJX6', 'RwXQZcECRu'], "first_words('strings1.txt')")

  def test6(self):
    res = first_words('wordlist1.txt')
    self.assertEqual(res, ['happy', 'sad', 'angsty', 'euphoric', 'maudlin'], "first_words('wordlist1.txt')")

myTests().main()
</tests>

      </program>
    </exercise>
    <exercise xml:id="exer-print_last_line-1" label="exe-sec-HW9-1-32179">
      <title>Print Last Line and Count of Lines</title>
      <statement>
        <p>The function <c>print_last_line</c> takes one parameter, <c>fname</c>, the name of a text file. The function should open the file for reading, read the lines of the file until it reaches the end, print out the last line in the file, report the number of lines contained in the file, and close the file.</p>
        <p><alert>Hint</alert>: Use a <alert>for</alert>loop to iterate over the lines and accumulate the count.</p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="60%" />
          <col top="minor" width="40%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>print_last_line("wordlist1.txt")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>maudlin

The file has 5 lines.</cd>
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <preamble visible="no">
_datafiles = {'cards.txt': 'Three of Clubs\nSix of Diamonds\nTen of Hearts\nEight of Clubs\nTwo of Clubs\nSix of Spades\nEight of Spades\nThree of Diamonds\nQueen of Diamonds\nTwo of Spades\nKing of Spades\nFive of Spades\nJack of Clubs\nFive of Hearts\nFour of Clubs\nNine of Clubs\nFive of Clubs\nTen of Spades\nThree of Spades\nJack of Diamonds\nNine of Spades\nSeven of Hearts\nEight of Diamonds\nSeven of Clubs\nFive of Diamonds\nQueen of Clubs\nJack of Hearts\nKing of Diamonds\nSix of Hearts\nFour of Diamonds\nAce of Hearts\nQueen of Hearts\nTwo of Diamonds\nKing of Clubs\nTwo of Hearts\nAce of Spades\nFour of Hearts\nFour of Spades\nNine of Diamonds\nSeven of Spades\nEight of Hearts\nSeven of Diamonds\nKing of Hearts\nSix of Clubs\nAce of Clubs\nNine of Hearts\nJack of Spades\nTen of Diamonds\nQueen of Spades\nThree of Hearts\nTen of Clubs\nAce of Diamonds\n', 'clocktimes.txt': '09:05\n09:15\n09:20\n09:30\n10:20\n10:25\n10:30\n10:40\n10:45\n11:00\n11:15\n11:30\n11:35\n11:55\n12:00\n12:05\n12:10\n12:45\n12:50\n12:55\n13:05\n13:20\n13:35\n14:00\n14:10\n14:15\n14:20\n14:25\n14:35\n14:45\n14:50\n15:00\n15:55\n16:00\n16:05\n16:20\n16:25\n16:30\n16:35\n16:40\n16:50\n16:55\n', 'jazzscales.txt': 'Play a major pent scale in Ab with attack on the 3rd string\nPlay a 7th pent scale in C# (Db) with attack on the 1st string\nPlay a major scale in B with attack on the 1st string\n', 'random1.txt': '77635\n51437\n25094\n13874\n24429\n88575\n90277\n72340\n8826\n36166\n99144\n6920\n56237\n58141\n77661\n51028\n93722\n14354\n26125\n10151\n75322\n', 'strings1.txt': '43XkyLl3QO\nKgYH4yZ9vp\nB9z0d27BWf\n3lVimQBNOl\nPZF1FkRCC0\naodmjRnbSM\nQinMgOKyJ3\nCIRpxzuBov\nbl2VB97vGe\nypinbN0yjw\nJotIY4A9GC\nwjlureOXLy\nDlHXQlUuhv\nQHvicuwzKp\nMzORTv2BYU\nvVnuElBS0s\naezB6hMOqj\nhipsag5OAM\nLhgNsYqWhB\n9b3yi6e29s\n6fx3eADW4n\nV28c1LIoQC\nIsN0xS6bZ4\nd1e0b2hJX6\nRwXQZcECRu\n', 'wordlist1.txt': 'happy\nsad\nangsty\neuphoric\nmaudlin\n'}

# Start of FileLike
class FileLike:
  def __init__(self, contents):
    self.contents = contents
    self.cursor = 0

  def __enter__(self):
    return self

  def __exit__(self, exc_type, exc_value, traceback):
    self.close()

  def tell(self):
    return self.cursor

  def seek(self, n):
    self.cursor = n

  def read(self):
    n = self.cursor
    self.cursor = max(n, len(self.contents))
    return self.contents[n:]

  def readline(self):
    n = self.cursor
    rest = self.contents[n:]
    index = rest.find("\n")
    if index == -1:
      self.cursor = len(self.contents)
      return rest
    index += 1
    self.cursor = n + index
    return rest[:index]

  def readlines(self):
    results = []
    while True:
      line = self.readline()
      if line == '':
        return results
      results.append(line)

  def __iter__(self):
    for l in self.readlines():
      yield l


  def close(self):
    pass

# End of FileLike

def open(fname, mode = "r"):
  """our own open"""
  if fname in _datafiles:
    return FileLike(_datafiles[fname])
  raise BaseException(f"File not found: {fname}")


        </preamble>
        <code>

        </code>
<tests>
from unittest.gui import TestCaseGui
import StringIO
import sys

class myTests(TestCaseGui):


  def test1(self):
    sys.stdout = StringIO.StringIO()
    print_last_line('cards.txt')
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    self.assertEqual(res, 'Ace of Diamonds\n\nThe file has 52 lines.\n', "print_last_line('cards.txt')")

  def test2(self):
    sys.stdout = StringIO.StringIO()
    print_last_line('clocktimes.txt')
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    self.assertEqual(res, '16:55\n\nThe file has 42 lines.\n', "print_last_line('clocktimes.txt')")

  def test3(self):
    sys.stdout = StringIO.StringIO()
    print_last_line('jazzscales.txt')
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    self.assertEqual(res, 'Play a major scale in B with attack on the 1st string\n\nThe file has 3 lines.\n', "print_last_line('jazzscales.txt')")

  def test4(self):
    sys.stdout = StringIO.StringIO()
    print_last_line('random1.txt')
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    self.assertEqual(res, '75322\n\nThe file has 21 lines.\n', "print_last_line('random1.txt')")

  def test5(self):
    sys.stdout = StringIO.StringIO()
    print_last_line('strings1.txt')
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    self.assertEqual(res, 'RwXQZcECRu\n\nThe file has 25 lines.\n', "print_last_line('strings1.txt')")

  def test6(self):
    sys.stdout = StringIO.StringIO()
    print_last_line('wordlist1.txt')
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    self.assertEqual(res, 'maudlin\n\nThe file has 5 lines.\n', "print_last_line('wordlist1.txt')")

myTests().main()
</tests>
      </program>
    </exercise>
    <exercise xml:id="exer-file_last_words-1" label="exe-sec-HW9-1-31853">
      <title>File Last Words</title>
      <statement>
        <p>The function <alert>last_words</alert>takes one parameter, <c>fname</c>, the name of a text file, and returns a list containing the last word from each line of the file. For example, if the file contents are:</p>
        <pre>apples are red
bananas are yellow
limes are green
        </pre>
        <p>then the list <c>["red", "yellow", "green"]</c> should be returned.</p>
        <p><alert>NOTE:</alert> You may assume the file will contain no blank lines.</p>
        <p><alert>BIG HINT:</alert> If <c>line</c>is a string representing a line of text (inside a for loop!), then <c>L = line.split()</c> creates a list of the words in the line.</p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="40%" />
          <col top="minor" width="60%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>last_words("wordlist1.txt")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>['happy', 'sad', 'angsty',
 'euphoric', 'maudlin']</cd>
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <preamble visible="no">
_datafiles = {'clocktimes.txt': '09:05\n09:15\n09:20\n09:30\n10:20\n10:25\n10:30\n10:40\n10:45\n11:00\n11:15\n11:30\n11:35\n11:55\n12:00\n12:05\n12:10\n12:45\n12:50\n12:55\n13:05\n13:20\n13:35\n14:00\n14:10\n14:15\n14:20\n14:25\n14:35\n14:45\n14:50\n15:00\n15:55\n16:00\n16:05\n16:20\n16:25\n16:30\n16:35\n16:40\n16:50\n16:55\n', 'jazzscales.txt': 'Play a major pent scale in Ab with attack on the 3rd string\nPlay a 7th pent scale in C# (Db) with attack on the 1st string\nPlay a major scale in B with attack on the 1st string\n', 'random1.txt': '77635\n51437\n25094\n13874\n24429\n88575\n90277\n72340\n8826\n36166\n99144\n6920\n56237\n58141\n77661\n51028\n93722\n14354\n26125\n10151\n75322\n', 'strings1.txt': '43XkyLl3QO\nKgYH4yZ9vp\nB9z0d27BWf\n3lVimQBNOl\nPZF1FkRCC0\naodmjRnbSM\nQinMgOKyJ3\nCIRpxzuBov\nbl2VB97vGe\nypinbN0yjw\nJotIY4A9GC\nwjlureOXLy\nDlHXQlUuhv\nQHvicuwzKp\nMzORTv2BYU\nvVnuElBS0s\naezB6hMOqj\nhipsag5OAM\nLhgNsYqWhB\n9b3yi6e29s\n6fx3eADW4n\nV28c1LIoQC\nIsN0xS6bZ4\nd1e0b2hJX6\nRwXQZcECRu\n', 'wordlist1.txt': 'happy\nsad\nangsty\neuphoric\nmaudlin\n', 'cards.txt': 'Three of Clubs\nSix of Diamonds\nTen of Hearts\nEight of Clubs\nTwo of Clubs\nSix of Spades\nEight of Spades\nThree of Diamonds\nQueen of Diamonds\nTwo of Spades\nKing of Spades\nFive of Spades\nJack of Clubs\nFive of Hearts\nFour of Clubs\nNine of Clubs\nFive of Clubs\nTen of Spades\nThree of Spades\nJack of Diamonds\nNine of Spades\nSeven of Hearts\nEight of Diamonds\nSeven of Clubs\nFive of Diamonds\nQueen of Clubs\nJack of Hearts\nKing of Diamonds\nSix of Hearts\nFour of Diamonds\nAce of Hearts\nQueen of Hearts\nTwo of Diamonds\nKing of Clubs\nTwo of Hearts\nAce of Spades\nFour of Hearts\nFour of Spades\nNine of Diamonds\nSeven of Spades\nEight of Hearts\nSeven of Diamonds\nKing of Hearts\nSix of Clubs\nAce of Clubs\nNine of Hearts\nJack of Spades\nTen of Diamonds\nQueen of Spades\nThree of Hearts\nTen of Clubs\nAce of Diamonds\n'}

# Start of FileLike
class FileLike:
  def __init__(self, contents):
    self.contents = contents
    self.cursor = 0

  def __enter__(self):
    return self

  def __exit__(self, exc_type, exc_value, traceback):
    self.close()

  def tell(self):
    return self.cursor

  def seek(self, n):
    self.cursor = n

  def read(self):
    n = self.cursor
    self.cursor = max(n, len(self.contents))
    return self.contents[n:]

  def readline(self):
    n = self.cursor
    rest = self.contents[n:]
    index = rest.find("\n")
    if index == -1:
      self.cursor = len(self.contents)
      return rest
    index += 1
    self.cursor = n + index
    return rest[:index]

  def readlines(self):
    results = []
    while True:
      line = self.readline()
      if line == '':
        return results
      results.append(line)

  def __iter__(self):
    for l in self.readlines():
      yield l


  def close(self):
    pass

# End of FileLike

def open(fname, mode = "r"):
  """our own open"""
  if fname in _datafiles:
    return FileLike(_datafiles[fname])
  raise BaseException(f"File not found: {fname}")

        </preamble>
        <code>

        </code>
        <postamble visible="no">
        </postamble>
<tests>
from unittest.gui import TestCaseGui
import StringIO
import sys

class myTests(TestCaseGui):


  def test1(self):
    res = last_words('cards.txt')
    self.assertEqual(res, ['Clubs', 'Diamonds', 'Hearts', 'Clubs', 'Clubs', 'Spades', 'Spades', 'Diamonds', 'Diamonds', 'Spades', 'Spades', 'Spades', 'Clubs', 'Hearts', 'Clubs', 'Clubs', 'Clubs', 'Spades', 'Spades', 'Diamonds', 'Spades', 'Hearts', 'Diamonds', 'Clubs', 'Diamonds', 'Clubs', 'Hearts', 'Diamonds', 'Hearts', 'Diamonds', 'Hearts', 'Hearts', 'Diamonds', 'Clubs', 'Hearts', 'Spades', 'Hearts', 'Spades', 'Diamonds', 'Spades', 'Hearts', 'Diamonds', 'Hearts', 'Clubs', 'Clubs', 'Hearts', 'Spades', 'Diamonds', 'Spades', 'Hearts', 'Clubs', 'Diamonds'], "last_words('cards.txt')")

  def test2(self):
    res = last_words('clocktimes.txt')
    self.assertEqual(res, ['09:05', '09:15', '09:20', '09:30', '10:20', '10:25', '10:30', '10:40', '10:45', '11:00', '11:15', '11:30', '11:35', '11:55', '12:00', '12:05', '12:10', '12:45', '12:50', '12:55', '13:05', '13:20', '13:35', '14:00', '14:10', '14:15', '14:20', '14:25', '14:35', '14:45', '14:50', '15:00', '15:55', '16:00', '16:05', '16:20', '16:25', '16:30', '16:35', '16:40', '16:50', '16:55'], "last_words('clocktimes.txt')")

  def test3(self):
    res = last_words('jazzscales.txt')
    self.assertEqual(res, ['string', 'string', 'string'], "last_words('jazzscales.txt')")

  def test4(self):
    res = last_words('random1.txt')
    self.assertEqual(res, ['77635', '51437', '25094', '13874', '24429', '88575', '90277', '72340', '8826', '36166', '99144', '6920', '56237', '58141', '77661', '51028', '93722', '14354', '26125', '10151', '75322'], "last_words('random1.txt')")

  def test5(self):
    res = last_words('strings1.txt')
    self.assertEqual(res, ['43XkyLl3QO', 'KgYH4yZ9vp', 'B9z0d27BWf', '3lVimQBNOl', 'PZF1FkRCC0', 'aodmjRnbSM', 'QinMgOKyJ3', 'CIRpxzuBov', 'bl2VB97vGe', 'ypinbN0yjw', 'JotIY4A9GC', 'wjlureOXLy', 'DlHXQlUuhv', 'QHvicuwzKp', 'MzORTv2BYU', 'vVnuElBS0s', 'aezB6hMOqj', 'hipsag5OAM', 'LhgNsYqWhB', '9b3yi6e29s', '6fx3eADW4n', 'V28c1LIoQC', 'IsN0xS6bZ4', 'd1e0b2hJX6', 'RwXQZcECRu'], "last_words('strings1.txt')")

  def test6(self):
    res = last_words('wordlist1.txt')
    self.assertEqual(res, ['happy', 'sad', 'angsty', 'euphoric', 'maudlin'], "last_words('wordlist1.txt')")

myTests().main()
</tests>
      </program>
    </exercise>
    <exercise xml:id="exer-file_print_line_if_exists-1" label="exe-sec-HW9-1-32166">
      <title>Print Line If It Exists</title>
      <statement>
        <p>The function <c>print_line_if_exists</c> takes two parameters, <c>fname</c>(a string, the name of a text file) and <c>line_num</c> (a non-negative integer, supposedly indexing a line in the file).</p>
        <p>
          <ul>
            <li>The function should open the file for reading, print the indicated line from the file, and close the file.</li>
            <li>However, if <c>line_num</c> is greater than or equal to the number of lines in the file, a warning message should printed instead.</li>
          </ul>
        </p>
        <p><alert>NOTE:</alert> <c>line_num</c> begins indexing at 0; line_num = 0 refers to the FIRST line of the file, line_num = 1 refers to the SECOND line, and so on.</p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="60%" />
          <col top="minor" width="40%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>print_line_if_exists("cards.txt", 51)</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>Ace of Diamonds</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>print_line_if_exists("cards.txt", 52)</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>Warning! line_num = 52 is not valid</cd>
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <preamble visible="no">
_datafiles = {'cards.txt': 'Three of Clubs\nSix of Diamonds\nTen of Hearts\nEight of Clubs\nTwo of Clubs\nSix of Spades\nEight of Spades\nThree of Diamonds\nQueen of Diamonds\nTwo of Spades\nKing of Spades\nFive of Spades\nJack of Clubs\nFive of Hearts\nFour of Clubs\nNine of Clubs\nFive of Clubs\nTen of Spades\nThree of Spades\nJack of Diamonds\nNine of Spades\nSeven of Hearts\nEight of Diamonds\nSeven of Clubs\nFive of Diamonds\nQueen of Clubs\nJack of Hearts\nKing of Diamonds\nSix of Hearts\nFour of Diamonds\nAce of Hearts\nQueen of Hearts\nTwo of Diamonds\nKing of Clubs\nTwo of Hearts\nAce of Spades\nFour of Hearts\nFour of Spades\nNine of Diamonds\nSeven of Spades\nEight of Hearts\nSeven of Diamonds\nKing of Hearts\nSix of Clubs\nAce of Clubs\nNine of Hearts\nJack of Spades\nTen of Diamonds\nQueen of Spades\nThree of Hearts\nTen of Clubs\nAce of Diamonds\n', 'strings1.txt': '43XkyLl3QO\nKgYH4yZ9vp\nB9z0d27BWf\n3lVimQBNOl\nPZF1FkRCC0\naodmjRnbSM\nQinMgOKyJ3\nCIRpxzuBov\nbl2VB97vGe\nypinbN0yjw\nJotIY4A9GC\nwjlureOXLy\nDlHXQlUuhv\nQHvicuwzKp\nMzORTv2BYU\nvVnuElBS0s\naezB6hMOqj\nhipsag5OAM\nLhgNsYqWhB\n9b3yi6e29s\n6fx3eADW4n\nV28c1LIoQC\nIsN0xS6bZ4\nd1e0b2hJX6\nRwXQZcECRu\n', 'wordlist1.txt': 'happy\nsad\nangsty\neuphoric\nmaudlin\n', 'random1.txt': '77635\n51437\n25094\n13874\n24429\n88575\n90277\n72340\n8826\n36166\n99144\n6920\n56237\n58141\n77661\n51028\n93722\n14354\n26125\n10151\n75322\n', 'clocktimes.txt': '09:05\n09:15\n09:20\n09:30\n10:20\n10:25\n10:30\n10:40\n10:45\n11:00\n11:15\n11:30\n11:35\n11:55\n12:00\n12:05\n12:10\n12:45\n12:50\n12:55\n13:05\n13:20\n13:35\n14:00\n14:10\n14:15\n14:20\n14:25\n14:35\n14:45\n14:50\n15:00\n15:55\n16:00\n16:05\n16:20\n16:25\n16:30\n16:35\n16:40\n16:50\n16:55\n', 'jazzscales.txt': 'Play a major pent scale in Ab with attack on the 3rd string\nPlay a 7th pent scale in C# (Db) with attack on the 1st string\nPlay a major scale in B with attack on the 1st string\n'}

# Start of FileLike
class FileLike:
  def __init__(self, contents):
    self.contents = contents
    self.cursor = 0

  def __enter__(self):
    return self

  def __exit__(self, exc_type, exc_value, traceback):
    self.close()

  def tell(self):
    return self.cursor

  def seek(self, n):
    self.cursor = n

  def read(self):
    n = self.cursor
    self.cursor = max(n, len(self.contents))
    return self.contents[n:]

  def readline(self):
    n = self.cursor
    rest = self.contents[n:]
    index = rest.find("\n")
    if index == -1:
      self.cursor = len(self.contents)
      return rest
    index += 1
    self.cursor = n + index
    return rest[:index]

  def readlines(self):
    results = []
    while True:
      line = self.readline()
      if line == '':
        return results
      results.append(line)

  def __iter__(self):
    for l in self.readlines():
      yield l


  def close(self):
    pass

# End of FileLike

def open(fname, mode = "r"):
  """our own open"""
  if fname in _datafiles:
    return FileLike(_datafiles[fname])
  raise BaseException(f"File not found: {fname}")

        </preamble>
        <code>

        </code>
<tests>
from unittest.gui import TestCaseGui
import StringIO
import sys

class myTests(TestCaseGui):


  def test1(self):
    sys.stdout = StringIO.StringIO()
    print_line_if_exists('wordlist1.txt', 5)
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    self.assertEqual(res, 'Warning! line_num = 5 is not valid\n', "print_line_if_exists('wordlist1.txt', 5)")

  def test2(self):
    sys.stdout = StringIO.StringIO()
    print_line_if_exists('random1.txt', 0)
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    self.assertEqual(res, '77635\n\n', "print_line_if_exists('random1.txt', 0)")

  def test3(self):
    sys.stdout = StringIO.StringIO()
    print_line_if_exists('jazzscales.txt', 12)
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    self.assertEqual(res, 'Warning! line_num = 12 is not valid\n', "print_line_if_exists('jazzscales.txt', 12)")

  def test4(self):
    sys.stdout = StringIO.StringIO()
    print_line_if_exists('clocktimes.txt', 7)
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    self.assertEqual(res, '10:40\n\n', "print_line_if_exists('clocktimes.txt', 7)")

  def test5(self):
    sys.stdout = StringIO.StringIO()
    print_line_if_exists('cards.txt', 51)
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    self.assertEqual(res, 'Ace of Diamonds\n\n', "print_line_if_exists('cards.txt', 51)")

  def test6(self):
    sys.stdout = StringIO.StringIO()
    print_line_if_exists('cards.txt', 52)
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    self.assertEqual(res, 'Warning! line_num = 52 is not valid\n', "print_line_if_exists('cards.txt', 52)")

  def test7(self):
    sys.stdout = StringIO.StringIO()
    print_line_if_exists('strings1.txt', 3)
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    self.assertEqual(res, '3lVimQBNOl\n\n', "print_line_if_exists('strings1.txt', 3)")

  def test8(self):
    sys.stdout = StringIO.StringIO()
    print_line_if_exists('strings1.txt', 3000)
    res = sys.stdout.getvalue()
    sys.stdout = sys.__stdout__
    self.assertEqual(res, 'Warning! line_num = 3000 is not valid\n', "print_line_if_exists('strings1.txt', 3000)")

myTests().main()
</tests>
      </program>
    </exercise>
  </exercises>
</section>
