<section xml:id="sec-HW5-1-practice">
  <title>HW5 Practice</title>
  <introduction />
  <exercises>
    <exercise xml:id="exer-slice_three_from_front-1-practice" label="exe-sec-HW5-1-31939-practice">
      <title>First Three Characters in a String</title>
      <statement>
        <p>The function <c>first_three</c> takes one string parameter, <alert>word</alert>, and <alert>returns</alert>a slice of length three from the front of word.  However, if word has length less than three, then the <c>None</c> object is returned instead.</p>
        <p>Hint: Use an if-else decision and the slicing operator [ : ] to solve this problem.</p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="40%" />
          <col top="minor" width="60%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>first_three("happy")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>hap</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>first_three("dog")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>dog</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>first_three("")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>None</cd>
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <code>

        </code>
        <tests>from unittest.gui import TestCaseGui
import StringIO
import sys

class myTests(TestCaseGui):


  def test1(self):
    res = first_three("happy")
    self.assertEqual(res, "hap", "first_three('happy')")

  def test2(self):
    res = first_three("NO")
    self.assertEqual(res, None, "first_three('NO')")

  def test3(self):
    res = first_three("dog")
    self.assertEqual(res, "dog", "first_three('dog')")

  def test4(self):
    res = first_three("")
    self.assertEqual(res, None, "first_three('')")

  def test5(self):
    res = first_three("Mississippi")
    self.assertEqual(res, "Mis", "first_three('Mississippi')")

  def test6(self):
    res = first_three("SPAM")
    self.assertEqual(res, "SPA", "first_three('SPAM')")

  def test7(self):
    res = first_three("!@#$%^*")
    self.assertEqual(res, "!@#", "first_three('!@#$%^*')")

  def test8(self):
    res = first_three("h")
    self.assertEqual(res, None, "first_three('h')")

myTests().main()</tests>
      </program>
    </exercise>
    <exercise xml:id="exer-slice_end-1-practice" label="exe-sec-HW5-1-31920-practice">
      <title />
      <statement>
        <h3>Slice the End of a String</h3>
        <p>The function <c>slice_end</c> takes a string parameter, <alert>word</alert>, and a positive integer <alert>n</alert>; the function <alert>returns</alert>a slice of length n from the end of word.  However, if word has length less than n, then the <alert>None</alert> object is returned instead.</p>
        <p>Hint: Notice that <em>returning the empty string when n is zero</em>might be a special case.</p>
        <p>You can use an if-elif-else decision structure and the slicing operator [ : ] to solve this problem.</p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="40%" />
          <col top="minor" width="60%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>slice_end("pineapple", 5)</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>apple</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>slice_end("too short", 20)</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>None</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>slice_end("llama", 0)</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd />
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>slice_end("Saturn", 6)</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>Saturn</cd>
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <code>

        </code>
        <tests>from unittest.gui import TestCaseGui
import StringIO
import sys

class myTests(TestCaseGui):


  def test1(self):
    res = slice_end("pineapple", 5)
    self.assertEqual(res, 'apple', 'slice_end("pineapple", 5)')

  def test2(self):
    res = slice_end("Enchilada", 7)
    self.assertEqual(res, 'chilada', 'slice_end("Enchilada", 7)')

  def test3(self):
    res = slice_end("too short", 20)
    self.assertEqual(res, None, 'slice_end("too short", 20)')

  def test4(self):
    res = slice_end("llama", 0)
    self.assertEqual(res, '', 'slice_end("llama", 0)')

  def test5(self):
    res = slice_end("Washington", 1)
    self.assertEqual(res, 'n', 'slice_end("Washington", 1)')

  def test6(self):
    res = slice_end("Saturn", 6)
    self.assertEqual(res, 'Saturn', 'slice_end("Saturn", 6)')

  def test7(self):
    res = slice_end("Hanover College", 7)
    self.assertEqual(res, 'College', 'slice_end("Hanover College", 7)')

myTests().main()</tests>
      </program>
    </exercise>
    <exercise xml:id="exer-slice_middle-1-practice" label="exe-sec-HW5-1-31924-practice">
      <title>Slice the Middle of a String</title>
      <statement>
        <p>The function <c>slice_middle</c> takes a string parameter, <alert><c>word</c>.</alert></p>
        <p>
          <ul>
            <li>If the length of word is less than 3, the function returns the <alert>empty</alert> string.</li>
            <li>Otherwise, the function returns a substring made of all but the first and last characters in word.</li>
          </ul>
        </p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="40%" />
          <col top="minor" width="60%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>slice_middle("onion")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>nio</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>slice_middle("?")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd />
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>slice_middle("slicing")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>licin</cd>
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <code>

        </code>
        <tests>from unittest.gui import TestCaseGui
import StringIO
import sys

class myTests(TestCaseGui):


  def test1(self):
    res = slice_middle('onion')
    self.assertEqual(res, 'nio', "slice_middle('onion')")

  def test2(self):
    res = slice_middle('dam')
    self.assertEqual(res, 'a', "slice_middle('dam')")

  def test3(self):
    res = slice_middle('?')
    self.assertEqual(res, '', "slice_middle('?')")

  def test4(self):
    res = slice_middle('PIZZA')
    self.assertEqual(res, 'IZZ', "slice_middle('PIZZA')")

  def test5(self):
    res = slice_middle('abracadabra')
    self.assertEqual(res, 'bracadabr', "slice_middle('abracadabra')")

  def test6(self):
    res = slice_middle('slicing')
    self.assertEqual(res, 'licin', "slice_middle('slicing')")

  def test7(self):
    res = slice_middle('')
    self.assertEqual(res, '', "slice_middle('')")

  def test8(self):
    res = slice_middle('factorial')
    self.assertEqual(res, 'actoria', "slice_middle('factorial')")

  def test9(self):
    res = slice_middle('pqr')
    self.assertEqual(res, 'q', "slice_middle('pqr')")

myTests().main()</tests>
      </program>
    </exercise>
    <exercise xml:id="exer-add_char_counts-1-practice" label="exe-sec-HW5-1-30807-practice">
      <title>Add Character Counts</title>
      <statement>
        <p>The function <c>add_char_counts</c> takes three parameters: a string <c>s</c> and two strings of length 1, <c>char1</c> and <c>char2</c>, and returns the total number of times either char1 or char2 appears in s.</p>
        <p>For example, the word "scallions" has 2 s's and 1 n, so <c>add_char_counts("scallions", "s", "n")</c> returns 3.</p>
        <p>Hint: Use a <alert>string method</alert> to help solve this problem (<c>count</c>).  You do not need an accumulator.</p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="80%" />
          <col top="minor" width="20%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>add_char_counts('Indiana', 'a', 'b')</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>2</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>add_char_counts('BAHAMAS', 'A', 'B')</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>4</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>add_char_counts('Mississippi', 'i', 'M')</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>5</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>add_char_counts("spam spam spam spam", "s", "p")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>8</cd>
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <code>

        </code>
        <tests>from unittest.gui import TestCaseGui
import StringIO
import sys

class myTests(TestCaseGui):


  def test1(self):
    res = add_char_counts('Indiana', 'a', 'b')
    self.assertEqual(res, 2, "add_char_counts('Indiana', 'a', 'b')")

  def test2(self):
    res = add_char_counts('BAHAMAS', 'A', 'B')
    self.assertEqual(res, 4, "add_char_counts('BAHAMAS', 'A', 'B')")

  def test3(self):
    res = add_char_counts('Mississippi', 'i', 'M')
    self.assertEqual(res, 5, "add_char_counts('Mississippi', 'i', 'M')")

  def test4(self):
    res = add_char_counts('pumpkin', 'e', 'g')
    self.assertEqual(res, 0, "add_char_counts('pumpkin', 'e', 'g')")

  def test5(self):
    res = add_char_counts('=:^.^:=', '=', '.')
    self.assertEqual(res, 3, "add_char_counts('=:^.^:=', '=', '.')")

  def test6(self):
    res = add_char_counts('spam spam spam spam', 's', 'p')
    self.assertEqual(res, 8, "add_char_counts('spam spam spam spam', 's', 'p')")

  def test7(self):
    res = add_char_counts('soupspoons', 's', 'o')
    self.assertEqual(res, 6, "add_char_counts('soupspoons', 's', 'o')")

  def test8(self):
    res = add_char_counts('abracadabra', 'a', 'a')
    self.assertEqual(res, 10, "add_char_counts('abracadabra', 'a', 'a')")

  def test9(self):
    res = add_char_counts('She sells sea shells down by the sea shore.', ' ', 's')
    self.assertEqual(res, 15, "add_char_counts('She sells sea shells down by the sea shore.', ' ', 's')")

myTests().main()</tests>
      </program>
    </exercise>
    <exercise xml:id="exer-count_words_of_length_ACCUM-1-practice" label="exe-sec-HW5-1-31943-practice">
      <title>Count Words of a Given Length</title>
      <statement>
        <p>The function <c>countWordsOfLength</c> takes a list of words L and an integer <c>n</c>, and returns the total number of <alert>words of length n</alert> that appear in the list. For example, <c>countWordsOfLength(["yes", "no", "maybe", "NNOOO"], 5)</c> would return 2 because the given list has two words of length 5.</p>
        <p>Hint: Use the <alert>accumulator</alert> pattern to solve this problem. Inside the for loop, use an <alert>if</alert>statement to only update the accumulator when the current word has length n.</p>
        <p>Hint: Instead of using a for loop with a <c>range</c>, you can iterate directly over  the elements of L.</p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="80%" />
          <col top="minor" width="20%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>L = ['a', 'b', 'c', 'de', 'fgh', 'i', 'jkl']
countWordsOfLength(L, 1)</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>4</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>L = ['a', 'b', 'c', 'de', 'fgh', 'i', 'jkl']
countWordsOfLength(L, 5)</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>0</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>L = ['bacon', 'onion', 'pepper', 'garlic', 'paprika', 'chipotle chili powder', 'salt', 'black pepper', 'vinegar', 'bbq sauce', 'ketchup', 'mustard', 'brown sugar', 'beans']
countWordsOfLength(L, 5)</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>3</cd>
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <code>

        </code>
        <tests>from unittest.gui import TestCaseGui
import StringIO
import sys

class myTests(TestCaseGui):


  def test1(self):
    res = countWordsOfLength(['a', 'b', 'c', 'de', 'fgh', 'i', 'jkl'], 1)
    self.assertEqual(res, 4, "countWordsOfLength(['a', 'b', 'c', 'de', 'fgh', 'i', 'jkl'], 1)")

  def test2(self):
    res = countWordsOfLength(['a', 'b', 'c', 'de', 'fgh', 'i', 'jkl'], 3)
    self.assertEqual(res, 2, "countWordsOfLength(['a', 'b', 'c', 'de', 'fgh', 'i', 'jkl'], 3)")

  def test3(self):
    res = countWordsOfLength(['a', 'b', 'c', 'de', 'fgh', 'i', 'jkl'], 5)
    self.assertEqual(res, 0, "countWordsOfLength(['a', 'b', 'c', 'de', 'fgh', 'i', 'jkl'], 5)")

  def test4(self):
    res = countWordsOfLength(['bacon', 'onion', 'pepper', 'garlic', 'paprika', 'chipotle chili powder', 'salt', 'black pepper', 'vinegar', 'bbq sauce', 'ketchup', 'mustard', 'brown sugar', 'beans'], 12)
    self.assertEqual(res, 1, "countWordsOfLength(['bacon', 'onion', 'pepper', 'garlic', 'paprika', 'chipotle chili powder', 'salt', 'black pepper', 'vinegar', 'bbq sauce', 'ketchup', 'mustard', 'brown sugar', 'beans'], 12)")

  def test5(self):
    res = countWordsOfLength(['bacon', 'onion', 'pepper', 'garlic', 'paprika', 'chipotle chili powder', 'salt', 'black pepper', 'vinegar', 'bbq sauce', 'ketchup', 'mustard', 'brown sugar', 'beans'], 7)
    self.assertEqual(res, 4, "countWordsOfLength(['bacon', 'onion', 'pepper', 'garlic', 'paprika', 'chipotle chili powder', 'salt', 'black pepper', 'vinegar', 'bbq sauce', 'ketchup', 'mustard', 'brown sugar', 'beans'], 7)")

  def test6(self):
    res = countWordsOfLength(['bacon', 'onion', 'pepper', 'garlic', 'paprika', 'chipotle chili powder', 'salt', 'black pepper', 'vinegar', 'bbq sauce', 'ketchup', 'mustard', 'brown sugar', 'beans'], 5)
    self.assertEqual(res, 3, "countWordsOfLength(['bacon', 'onion', 'pepper', 'garlic', 'paprika', 'chipotle chili powder', 'salt', 'black pepper', 'vinegar', 'bbq sauce', 'ketchup', 'mustard', 'brown sugar', 'beans'], 5)")

  def test7(self):
    res = countWordsOfLength(['bacon', 'onion', 'pepper', 'garlic', 'paprika', 'chipotle chili powder', 'salt', 'black pepper', 'vinegar', 'bbq sauce', 'ketchup', 'mustard', 'brown sugar', 'beans'], 0)
    self.assertEqual(res, 0, "countWordsOfLength(['bacon', 'onion', 'pepper', 'garlic', 'paprika', 'chipotle chili powder', 'salt', 'black pepper', 'vinegar', 'bbq sauce', 'ketchup', 'mustard', 'brown sugar', 'beans'], 0)")

  def test8(self):
    res = countWordsOfLength([], 5)
    self.assertEqual(res, 0, "countWordsOfLength([], 5)")

  def test9(self):
    res = countWordsOfLength(['just one thing here'], 5)
    self.assertEqual(res, 0, "countWordsOfLength(['just one thing here'], 5)")

myTests().main()</tests>
      </program>
    </exercise>
    <exercise xml:id="exer-make_abbreviation-1-practice" label="exe-sec-HW5-1-31995-practice">
      <title>Make Abbreviation</title>
      <statement>
        <p>The function <c>make_abbreviation</c> takes one parameter, <c>wordlist</c>, a list of words, and produces an acronym out of the words in the list, with a period after each initial character.  The function returns the resulting acronym string.</p>
        <p>For example, the input <c>["International", "Business", "Machines"]</c> would cause the function to return <c>"I.B.M."</c></p>
        <p>Write a function definition for <c>make_abbreviation</c>.</p>
        <p>Hint: Use the accumulate-a-string pattern.</p>
        <p>- Initialize the accumulator variable to <c>""</c>, the empty string.</p>
        <p>- Iterate w over the items <c>wordlist</c></p>
        <p>- Inside the <alert>for</alert> loop, update the accumulator variable by concatenating it with the first character in the current word as well as a period.</p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="70%" />
          <col top="minor" width="30%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>wordlist = ["Post", "Meridiem"]
make_abbreviation(wordlist)</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>P.M.</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>wordlist = ["Joint", "Photographic", "Experts", "Group"]
make_abbreviation(wordlist)</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>J.P.E.G.</cd>
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <code>

        </code>
        <tests>
from unittest.gui import TestCaseGui
import StringIO
import sys

class myTests(TestCaseGui):

  def test1(self):
    res = make_abbreviation(['Post', 'Meridiem'])
    self.assertEqual(res, 'P.M.', "make_abbreviation(['Post', 'Meridiem'])")

  def test2(self):
    res = make_abbreviation(['Joint', 'Photographic', 'Experts', 'Group'])
    self.assertEqual(res, 'J.P.E.G.', "make_abbreviation(['Joint', 'Photographic', 'Experts', 'Group'])")

  def test3(self):
    res = make_abbreviation(['Subscriber', 'Identification', 'Module'])
    self.assertEqual(res, 'S.I.M.', "make_abbreviation(['Subscriber', 'Identification', 'Module'])")

  def test4(self):
    res = make_abbreviation(['What', 'You', 'See', 'Is', 'What', 'You', 'Get'])
    self.assertEqual(res, 'W.Y.S.I.W.Y.G.', "make_abbreviation(['What', 'You', 'See', 'Is', 'What', 'You', 'Get'])")

  def test5(self):
    res = make_abbreviation(['Monsieur'])
    self.assertEqual(res, 'M.', "make_abbreviation(['Monsieur'])")

  def test6(self):
    res = make_abbreviation(['Vital', 'Information', 'Resources', 'Under', 'Siege'])
    self.assertEqual(res, 'V.I.R.U.S.', "make_abbreviation(['Vital', 'Information', 'Resources', 'Under', 'Siege'])")

  def test7(self):
    res = make_abbreviation(['Wonderful', 'Exciting', 'Interesting', 'Real', 'Different'])
    self.assertEqual(res, 'W.E.I.R.D.', "make_abbreviation(['Wonderful', 'Exciting', 'Interesting', 'Real', 'Different'])")

  def test8(self):
    res = make_abbreviation(['Bubble', 'Aromatic', 'Therapeutic', 'Hot'])
    self.assertEqual(res, 'B.A.T.H.', "make_abbreviation(['Bubble', 'Aromatic', 'Therapeutic', 'Hot'])")

myTests().main()
</tests>      </program>
    </exercise>
    <exercise xml:id="exer-accum_count_contains_x-1-practice" label="exe-sec-HW5-1-161050-practice">
      <title>Count Contains 'X' or 'x'</title>
      <statement>
        <p>The function <c>count_contains_x</c> takes a list of strings (<c>wordList</c>) and <alert>returns</alert>the count of how many strings in the list contain at least one upper- or lower-case 'x'.</p>
        <p>Hint: Use the accumulator pattern to accumulate the count.</p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="80%" />
          <col top="minor" width="20%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>wordList = ["xerox", "OXEN", "whiffleball", "XOXOX"]
count_contains_x(wordList)</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>3</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>fish = ["angel", "gold", "cat", "puffer", "damsel"]
count_contains_x(fish)</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>0</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>wordList = ["X", "x"]
count_contains_x(wordList)</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>2</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>list_of_one_x = ["x"]
count_contains_x(list_of_one_x)</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>1</cd>
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <code>

        </code>
<tests>
from unittest.gui import TestCaseGui
import StringIO
import sys

class myTests(TestCaseGui):


  def test1(self):
    res = count_contains_x(['xerox', 'OXEN', 'whiffleball', 'XOXOX'])
    self.assertEqual(res, 3, "count_contains_x(['xerox', 'OXEN', 'whiffleball', 'XOXOX'])")

  def test2(self):
    res = count_contains_x(['angel', 'gold', 'cat', 'puffer', 'damsel'])
    self.assertEqual(res, 0, "count_contains_x(['angel', 'gold', 'cat', 'puffer', 'damsel'])")

  def test3(self):
    res = count_contains_x(['X', 'x'])
    self.assertEqual(res, 2, "count_contains_x(['X', 'x'])")

  def test4(self):
    res = count_contains_x(['OXYGEN'])
    self.assertEqual(res, 1, "count_contains_x(['OXYGEN'])")

  def test5(self):
    res = count_contains_x(['Box drum', 'Paixiao', 'Saxophone', 'Txistu', 'Xalam', 'Xaphoon', 'Xun', 'Xylophone', 'Bari Sax'])
    self.assertEqual(res, 9, "count_contains_x(['Box drum', 'Paixiao', 'Saxophone', 'Txistu', 'Xalam', 'Xaphoon', 'Xun', 'Xylophone', 'Bari Sax'])")

  def test6(self):
    res = count_contains_x(['atari', 'x-box', 'nintendo'])
    self.assertEqual(res, 1, "count_contains_x(['atari', 'x-box', 'nintendo'])")

  def test7(self):
    res = count_contains_x(['pixxxel', 'Xanadu', 'pEnGuInZ', 'piXXa', 'tomatox', 'alpaca', 'Matrix'])
    self.assertEqual(res, 5, "count_contains_x(['pixxxel', 'Xanadu', 'pEnGuInZ', 'piXXa', 'tomatox', 'alpaca', 'Matrix'])")

myTests().main()
</tests>      </program>
    </exercise>
    <exercise xml:id="exer-accum_remove_vowels-1-practice" label="exe-sec-HW5-1-31997-practice">
      <title>Remove Vowels</title>
      <statement>
        <p>The function <c>remove_vowels</c> takes one parameter, <c>s</c>, a string of any length. The function then returns a new string that contains all the letters in <c>s</c> except for the vowels: 'a', 'e', 'i', 'o', 'u'.</p>
        <p>You may assume the string contains no upper-case letters.</p>
        <p><alert>Hint</alert>: Accumulate the new string.
          <ul>
          <li><p>Use a for loop to iterate over the characters in s</p></li>
          <li><p>Inside the loop, update the accumulator if the current character is not a lowercase vowel ( <c>if ch not in 'aeiou'</c> )</p></li>
        </ul>
        For example:</p>
        <tabular>
          <col top="minor" width="40%" />
          <col top="minor" width="60%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>remove_vowels("giraffe")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>grff</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>remove_vowels("love bug")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>lv bg</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>remove_vowels("racecar")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>rccr</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>remove_vowels("uoiea")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd />
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <code>

        </code>
<tests>
from unittest.gui import TestCaseGui
import StringIO
import sys

class myTests(TestCaseGui):


  def test1(self):
    res = remove_vowels('giraffe')
    self.assertEqual(res, 'grff', "remove_vowels('giraffe')")

  def test2(self):
    res = remove_vowels('love bug')
    self.assertEqual(res, 'lv bg', "remove_vowels('love bug')")

  def test3(self):
    res = remove_vowels('racecar')
    self.assertEqual(res, 'rccr', "remove_vowels('racecar')")

  def test4(self):
    res = remove_vowels('iota')
    self.assertEqual(res, 't', "remove_vowels('iota')")

  def test5(self):
    res = remove_vowels('queue')
    self.assertEqual(res, 'q', "remove_vowels('queue')")

  def test6(self):
    res = remove_vowels('uoiea')
    self.assertEqual(res, '', "remove_vowels('uoiea')")

  def test7(self):
    res = remove_vowels('alpacas are cute')
    self.assertEqual(res, 'lpcs r ct', "remove_vowels('alpacas are cute')")

  def test8(self):
    res = remove_vowels('gypsy moth')
    self.assertEqual(res, 'gypsy mth', "remove_vowels('gypsy moth')")

  def test9(self):
    res = remove_vowels('cheetah')
    self.assertEqual(res, 'chth', "remove_vowels('cheetah')")

  def test10(self):
    res = remove_vowels('aieeioouuaiieeaiaoox')
    self.assertEqual(res, 'x', "remove_vowels('aieeioouuaiieeaiaoox')")

myTests().main()
</tests>      </program>
    </exercise>
    <exercise xml:id="exer-decode_ascii_ACCUM-1-practice" label="exe-sec-HW5-1-30808-practice">
      <title>Decode from ASCII</title>
      <statement>
        <p>The function <c>decodeFromAscii</c> takes a list <c>L</c> of integers in the range 32 through 126. Each item in L is an ASCII code representing a single printable character. The function accumulates the string of characters represented by the list of ASCII codes and returns that string.</p>
        <p>For example, <c>decodeFromAscii([78, 105, 99, 101, 33])</c> returns <c>"Nice!"</c> because N is chr(78), i is chr(105), c is chr(99), etc.</p>
        <p>Note: ASCII was created in the 1960's so programmers could have a standard encoding for common keyboard characters as numbers. It was later expanded to create Unicode so that characters from languages  other than English, mathematical symbols, and many other symbols could also be represented by numerical codes.</p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="80%" />
          <col top="minor" width="20%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>decodeFromAscii([36, 49, 44, 48, 48, 48])</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>$1,000</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>decodeFromAscii([115, 110, 111, 119])</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>snow</cd>
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <code>

        </code>
<tests>
from unittest.gui import TestCaseGui
import StringIO
import sys

class myTests(TestCaseGui):


  def test1(self):
    res = decodeFromAscii([36, 49, 44, 48, 48, 48])
    self.assertEqual(res, '$1,000', "decodeFromAscii([36, 49, 44, 48, 48, 48])")

  def test2(self):
    res = decodeFromAscii([115, 110, 111, 119])
    self.assertEqual(res, 'snow', "decodeFromAscii([115, 110, 111, 119])")

  def test3(self):
    res = decodeFromAscii([68, 79, 71, 83])
    self.assertEqual(res, 'DOGS', "decodeFromAscii([68, 79, 71, 83])")

  def test4(self):
    res = decodeFromAscii([80, 121, 116, 104, 111, 110])
    self.assertEqual(res, 'Python', "decodeFromAscii([80, 121, 116, 104, 111, 110])")

  def test5(self):
    res = decodeFromAscii([65, 115, 107, 32, 113, 117, 101, 115, 116, 105, 111, 110, 115, 46])
    self.assertEqual(res, 'Ask questions.', "decodeFromAscii([65, 115, 107, 32, 113, 117, 101, 115, 116, 105, 111, 110, 115, 46])")

  def test6(self):
    res = decodeFromAscii([84, 104, 105, 115, 32, 111, 110, 101, 32, 105, 115, 32, 115, 101, 99, 114, 101, 116, 33, 33, 33])
    self.assertEqual(res, 'This one is secret!!!', "decodeFromAscii([84, 104, 105, 115, 32, 111, 110, 101, 32, 105, 115, 32, 115, 101, 99, 114, 101, 116, 33, 33, 33])")

myTests().main()
</tests>      </program>
    </exercise>
    <exercise xml:id="exer-concat_alternate_characters_ACCUM-1-practice" label="exe-sec-HW5-1-30810-practice">
      <title>Concatenate Alternating Characters</title>
      <statement>
        <p>The function <c>concatAltChars</c> takes a string parameter <c>s</c> and a boolean parameter <c>even</c>; the function concatenates together either the even-position characters of s or the odd-position characters of s, depending on whether even is True, or False, respectively. The accumulated string is returned.</p>
        <p>For example, the even-indexed characters of "HANOVER" are at indexes 0, 2, 4, and 6:  H N V R. So, <c>concatAltChars("HANOVER", True)</c>returns the string result <c>"HNVR"</c>.</p>
        <p>But <c>concatAltChars("HANOVER", False)</c> returns the string result <c>"AOE"</c>.</p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="80%" />
          <col top="minor" width="20%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>concatAltChars('goldfinches', True)</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>glfnhs</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>concatAltChars("A", True)</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>A</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>concatAltChars("A", False)</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd />
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <code>

        </code>
<tests>
from unittest.gui import TestCaseGui
import StringIO
import sys

class myTests(TestCaseGui):


  def test1(self):
    res = concatAltChars('goldfinches', True)
    self.assertEqual(res, 'glfnhs', "concatAltChars('goldfinches', True)")

  def test2(self):
    res = concatAltChars('A', True)
    self.assertEqual(res, 'A', "concatAltChars('A', True)")

  def test3(self):
    res = concatAltChars('A', False)
    self.assertEqual(res, '', "concatAltChars('A', False)")

  def test4(self):
    res = concatAltChars('some popular animals on Google search are shark, bear, and horse', True)
    self.assertEqual(res, 'sm oua nml nGol erhaesak er n os', "concatAltChars('some popular animals on Google search are shark, bear, and horse', True)")

  def test5(self):
    res = concatAltChars('some popular animals on Google search are shark, bear, and horse', False)
    self.assertEqual(res, 'oepplraiaso ogesac r hr,ba,adhre', "concatAltChars('some popular animals on Google search are shark, bear, and horse', False)")

  def test6(self):
    res = concatAltChars('some popular animals on Google search are not bunny, worm, cow, sheep', True)
    self.assertEqual(res, 'sm oua nml nGol erhaentbny om o,sep', "concatAltChars('some popular animals on Google search are not bunny, worm, cow, sheep', True)")

  def test7(self):
    res = concatAltChars('some popular animals on Google search are not bunny, worm, cow, sheep', False)
    self.assertEqual(res, 'oepplraiaso ogesac r o un,wr,cw he', "concatAltChars('some popular animals on Google search are not bunny, worm, cow, sheep', False)")

myTests().main()
</tests>      </program>
    </exercise>
    <exercise xml:id="exer-string_has_repeat-1-practice" label="exe-sec-HW5-1-31054-practice">
      <title>String Has Repeated Character?</title>
      <statement>
        <p>Write the code to define a function <c>has_repeat(mystr)</c>, which takes a string parameter and returns a boolean result.</p>
        <p>- If mystr has a character that is repeated (<c>count</c> is more than 1 for that character), return <alert>True</alert></p>
        <p>- Otherwise, return <alert>False</alert></p>
        <p>Hint: You can solve this problem by iterating over the characters in mystr and comparing each character's <c>count</c> with 1.</p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="40%" />
          <col top="minor" width="60%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>has_repeat("Happy")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>True</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>has_repeat("sad")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>False</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>has_repeat("salmons")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>True</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>has_repeat("Trouts")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>False</cd>
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <code>

        </code>
<tests>
from unittest.gui import TestCaseGui
import StringIO
import sys

class myTests(TestCaseGui):


  def test1(self):
    res = has_repeat('Happy')
    self.assertEqual(res, True, "has_repeat('Happy')")

  def test2(self):
    res = has_repeat('sad')
    self.assertEqual(res, False, "has_repeat('sad')")

  def test3(self):
    res = has_repeat('salmons')
    self.assertEqual(res, True, "has_repeat('salmons')")

  def test4(self):
    res = has_repeat('Trouts')
    self.assertEqual(res, False, "has_repeat('Trouts')")

  def test5(self):
    res = has_repeat('3.14159')
    self.assertEqual(res, True, "has_repeat('3.14159')")

  def test6(self):
    res = has_repeat('dolphin')
    self.assertEqual(res, False, "has_repeat('dolphin')")

  def test7(self):
    res = has_repeat('')
    self.assertEqual(res, False, "has_repeat('')")

  def test8(self):
    res = has_repeat('1word1')
    self.assertEqual(res, True, "has_repeat('1word1')")

  def test9(self):
    res = has_repeat('1word2')
    self.assertEqual(res, False, "has_repeat('1word2')")

myTests().main()
</tests>      </program>
    </exercise>
    <exercise xml:id="exer-string_has_adjacent_repeats-1-practice" label="exe-sec-HW5-1-31055-practice">
      <title>String Has Adjacent Repeated Characters?</title>
      <statement>
        <p>Write the code to define the function <c>has_adjacent_repeats(mystr)</c>, which takes a string parameter and returns a boolean result.</p>
        <p>- If mystr has at least one instance of adjacent characters being equal, return True</p>
        <p>- Otherwise, return False</p>
        <p>Hint: You can iterate over the positive indices i for characters in mystr: 1, 2, 3, ...., len(mystr)-1 with a <alert>for</alert> loop. For each i, if the character at index i matches the character at index i - 1, return <alert>True</alert>. Otherwise, if no doubled character is found in the entire string, return <alert>False</alert>.</p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="40%" />
          <col top="minor" width="60%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>has_adjacent_repeats("NOODLES")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>True</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>has_adjacent_repeats("Bananas")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>False</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>has_adjacent_repeats("Hanoverr")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>True</cd>
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <code>

        </code>
<tests>
from unittest.gui import TestCaseGui
import StringIO
import sys

class myTests(TestCaseGui):

  def test1(self):
    res = has_adjacent_repeats('NOODLES')
    self.assertEqual(res, True, "has_adjacent_repeats('NOODLES')")

  def test2(self):
    res = has_adjacent_repeats('Bananas')
    self.assertEqual(res, False, "has_adjacent_repeats('Bananas')")

  def test3(self):
    res = has_adjacent_repeats('Hanoverr')
    self.assertEqual(res, True, "has_adjacent_repeats('Hanoverr')")

  def test4(self):
    res = has_adjacent_repeats('LLAMA')
    self.assertEqual(res, True, "has_adjacent_repeats('LLAMA')")

  def test5(self):
    res = has_adjacent_repeats('information')
    self.assertEqual(res, False, "has_adjacent_repeats('information')")

  def test6(self):
    res = has_adjacent_repeats('Hanover')
    self.assertEqual(res, False, "has_adjacent_repeats('Hanover')")

  def test7(self):
    res = has_adjacent_repeats('')
    self.assertEqual(res, False, "has_adjacent_repeats('')")

  def test8(self):
    res = has_adjacent_repeats('H')
    self.assertEqual(res, False, "has_adjacent_repeats('H')")

  def test9(self):
    res = has_adjacent_repeats('AA')
    self.assertEqual(res, True, "has_adjacent_repeats('AA')")

  def test10(self):
    res = has_adjacent_repeats('123456788')
    self.assertEqual(res, True, "has_adjacent_repeats('123456788')")

  def test11(self):
    res = has_adjacent_repeats('1x2x3x4x5x7x8x8')
    self.assertEqual(res, False, "has_adjacent_repeats('1x2x3x4x5x7x8x8')")

myTests().main()
</tests>      </program>
    </exercise>
    <exercise xml:id="exer-string_is_numberlike-1-practice" label="exe-sec-HW5-1-31056-practice">
      <title>String is Numberlike?</title>
      <statement>
        <p><c>mystr</c> is "numberlike" if:</p>
        <p>- <alert>mystr</alert> is not the empty string</p>
        <p>- the first character in <alert>mystr</alert> is a digit</p>
        <p>- the count of '.' in <alert>mystr</alert> is zero or one (can't have more than one decimal point in a number)</p>
        <p>- every character in <alert>mystr</alert> is a digit ('0', '1', '2', ..., '9') or a decimal point ('.')</p>
        <p>Write the code to define the function <c>is_numberlike(mystr)</c>, which takes a string parameter and returns a boolean result.</p>
        <p>If mystr is numberlike, the function returns <alert>True</alert>; otherwise, the function returns <alert>False</alert>.</p>
        <p>Hint: Check the first three requirements with decision statements, then use a <alert>for</alert>loop to iterate over mystr and check the every character is a digit or a dot</p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="80%" />
          <col top="minor" width="20%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>if not (is_numberlike("13.5") is True):
    print("error")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd />
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>if not (is_numberlike("-5532") is False):
    print("error")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd />
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>if not (is_numberlike("175") is True):
    print("error")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd />
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <code>
from string import digits

def is_numberlike(mystr):
    # Case 1: The empty string is not numberlike
    if mystr == "":
        return False

    ## more code here ##
        </code>
<tests>
from unittest.gui import TestCaseGui
import StringIO
import sys

class myTests(TestCaseGui):


  def test1(self):
    res = is_numberlike('13.5')
    self.assertEqual(res, True, "is_numberlike('13.5')")

  def test2(self):
    res = is_numberlike('-5532')
    self.assertEqual(res, False, "is_numberlike('-5532')")

  def test3(self):
    res = is_numberlike('175')
    self.assertEqual(res, True, "is_numberlike('175')")

  def test4(self):
    res = is_numberlike('3')
    self.assertEqual(res, True, "is_numberlike('3')")

  def test5(self):
    res = is_numberlike('.')
    self.assertEqual(res, False, "is_numberlike('.')")

  def test6(self):
    res = is_numberlike('')
    self.assertEqual(res, False, "is_numberlike('')")

  def test7(self):
    res = is_numberlike('3.141.59')
    self.assertEqual(res, False, "is_numberlike('3.141.59')")

  def test8(self):
    res = is_numberlike('88.59!')
    self.assertEqual(res, False, "is_numberlike('88.59!')")

  def test9(self):
    res = is_numberlike('1234B')
    self.assertEqual(res, False, "is_numberlike('1234B')")

  def test10(self):
    res = is_numberlike('1242498792472.9749247')
    self.assertEqual(res, True, "is_numberlike('1242498792472.9749247')")

  def test11(self):
    res = is_numberlike('3344.')
    self.assertEqual(res, True, "is_numberlike('3344.')")

  def test12(self):
    res = is_numberlike('332.2')
    self.assertEqual(res, True, "is_numberlike('332.2')")

  def test13(self):
    res = is_numberlike('54545')
    self.assertEqual(res, True, "is_numberlike('54545')")

  def test14(self):
    res = is_numberlike('0.')
    self.assertEqual(res, True, "is_numberlike('0.')")

  def test15(self):
    res = is_numberlike('1.2.3')
    self.assertEqual(res, False, "is_numberlike('1.2.3')")

myTests().main()
</tests>      </program>
    </exercise>
    <exercise xml:id="exer-find_max_char-1-practice" label="exe-sec-HW5-1-31988-practice">
      <title>Find Max Char</title>
      <statement>
        <p>The function <c>find_max_char</c> takes one string parameter, <c>mystr</c>. The function returns the character in <c>mystr</c> with the maximum ordinal value (the character that comes last in the ASCII alphabet, relative to the other characters in <c>mystr</c>). However, if <alert>mystr</alert> is the empty string, the function returns the Python object <alert>None</alert>.</p>
        <p>Notice that we don't need the ord() function to compare ordinal values; we can use comparison operators. For example, "Z" &gt; "C" is true because Z comes after C in the ASCII alphabet.</p>
        <p>We need to "accumulate" the correct answer (max_char) and then return it:</p>
        <p>
          <ul>
            <li>First, if mystr is empty, return None</li>
            <li>Then initialize max_char (accumulator variable) to the first character in mystr</li>
            <li>Use a for loop to iterate ch over the characters in mystr</li>
            <li>Each time a new maximum character is found (<c>if ch &gt; max_char</c>), max_char is updated to "remember" that the new maximum is ch</li>
            <li>Finally, return max_char</li>
          </ul>
        </p>
        <p>For example:</p>
        <tabular>
          <col top="minor" width="40%" />
          <col top="minor" width="60%" />
          <row header="yes" left="minor">
            <cell bottom="minor" right="minor">Test</cell>
            <cell bottom="minor" right="minor">Result</cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>find_max_char("naproxen")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>x</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>find_max_char("Gorilla Glue")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>u</cd>
              </p>
            </cell>
          </row>
          <row header="no" left="minor">
            <cell bottom="minor" right="minor">
              <p>
                <cd>find_max_char("")</cd>
              </p>
            </cell>
            <cell bottom="minor" right="minor">
              <p>
                <cd>None</cd>
              </p>
            </cell>
          </row>
        </tabular>
      </statement>
      <program interactive="activecode" language="python">
        <code>

        </code>
<tests>
from unittest.gui import TestCaseGui
import StringIO
import sys

class myTests(TestCaseGui):


  def test1(self):
    res = find_max_char('naproxen')
    self.assertEqual(res, 'x', "find_max_char('naproxen')")

  def test2(self):
    res = find_max_char('Gorilla Glue')
    self.assertEqual(res, 'u', "find_max_char('Gorilla Glue')")

  def test3(self):
    res = find_max_char('')
    self.assertEqual(res, None, "find_max_char('')")

  def test4(self):
    res = find_max_char('X')
    self.assertEqual(res, 'X', "find_max_char('X')")

  def test5(self):
    res = find_max_char('##$@#@%^$')
    self.assertEqual(res, '^', "find_max_char('##$@#@%^$')")

  def test6(self):
    res = find_max_char('zebras')
    self.assertEqual(res, 'z', "find_max_char('zebras')")

  def test7(self):
    res = find_max_char('7905')
    self.assertEqual(res, '9', "find_max_char('7905')")

  def test8(self):
    res = find_max_char('Hanover')
    self.assertEqual(res, 'v', "find_max_char('Hanover')")

myTests().main()
</tests>      </program>
    </exercise>
  </exercises>
</section>
